# Homework 3: More String, Lists, and Python Objects (18 pts)

name: Robert Frazier

email: robertcf@umich.edu

This homework assignment took me XX hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I discussed the final question with Will Meynardie where we discussed what conditons it exactly wants for a string to be valid.

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

```python
import collections #when should I have used this package?
import string
import itertools
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

```python
def primes():
    h = 2
    while True:
        is_prime = True
        for i in range(2, int(h**0.5) + 1): 
            if h % i == 0:
                is_prime = False
        if is_prime:
            yield h
        h += 1
```

```python
p = primes()
for i in range(5000):
    (next(p))
print(next(p))
```

Output:
```
48619
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

<!-- grade_id: ruler, points: 0, grade: false, solution: true -->
```python
def ruler():
    n = 1 #number i am starting on 
    reps = 1 #initial number of repetitions 
    adds = 0 #how many times ive added to n in a repeititon 
    while True:
        for i in range(reps):
            yield next(itertools.repeat(n,reps)) #for the number of reps yield n 
        n += 1 #now i increase my n 
        adds += 1 #and so if i increase my n that means adds has gone up 
        if adds == reps: #if adds equals the number of repeitions that means the numbers i am concerned about has appeared rep number of times
            reps += 1 #increase the number of repeitions
            adds = 0 #reset adds

#i am not very familiar with itertools and it seems like there's a lot of functions so im curious about what would be the best thing to do for this
#because i dont think what I did saved me much space
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
def look_say():
    line = "1"
    while True:
        if line == "1": #this part is jank but I just need a way to first spit out "1"
            yield(line)
        split = [] #splitting the word into groups based on repeating numbers
        search = line[0] #first number im searching for
        grab = [line[0]] #grab gets numbers added to it and then appended as a list to the list of lists that is split 
        for i in line[1:]: #going through the string besides the first digit which is specified as initial search
            if i == search:  #if I find that search repeats in succession 
                grab.append(i) #I add that digit to grab
            else:  #if the number following search isn't also search I do this 
                split.append(grab) #I add what's in grab so far to split 
                search = i #I switch search to the number I just landed on 
                grab = [i] #and I put this number into grab
        split.append(grab) #once i go through all the numbers the list will run out and so whatever's left in grab has to be added to split 
        result = [] #i now have to say what I see
        for i in split: #split will have lists that have lists of numbers 
            result.append(str(len(i))) #i will first say how many of a number is in a list in split 
            result.append(i[0]) #then the first number in the list in split is what that number is 
        result = "".join(result) #turn it into a nice string 
        line = result #change line to what i'll be outputting for next iteration
        yield(result)
```

```python
s = look_say()
for i in range(10):
    print(next(s))
```

Output:
```
1
11
21
1211
111221
312211
13112221
1113213211
31131211131221
13211311123113112211
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
    hold = list(s)
    letters = string.ascii_lowercase
    
    while k > 26: 
       k -= 26
    while k < -26:
        k += 26
    k = int(k)
    for i in range(len(s)):
        index = letters.find(s[i])
        if index + k < 26:
            hold[i] = letters[index+k]
        else:
            store = index+k - 26
            hold[i] = letters[store]
    return "".join(hold)
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    while k > 26: 
       k -= 26
    while k < -26:
        k += 26
    k = int(k)
    hold = list(s)
    letters = string.ascii_lowercase
    for i in range(len(s)):
        index = letters.find(s[i])
        if index - k < 26:
            hold[i] = letters[index-k]
        else:
            store = index-k-26
            hold[i] = letters[store]
    return "".join(hold)
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


```python
def is_valid(s):
    stack = []
    bracket_map = {')': '(', '}': '{', ']': '['}
    
    for char in s:
      #  print("Stack is", stack)
        if char in bracket_map:
            if stack:
                top_element = stack.pop()
            else:
                top_element = "Fiddlesticks"
            if bracket_map[char] != top_element:
                return False
        else:
            stack.append(char)
    return not stack
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```

```python

```