### Lecture 3-2 Practice: Lambda Functions and Dictionaries

### Question 1
Sorting with lambda key functions.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
# Sort list of numbers by absolute value
nums = [-4, 1, -2, 3]
# YOUR CODE HERE
assert(nums == [1, -2, 3, -4])

# Sort list of strings by length.
words = ['apple', 'cat', 'banana'] 
# YOUR CODE HERE
assert(words == ['cat', 'apple', 'banana'])
```

### Question 2
Given an array of intervals where `intervals[i] = [starti, endi]`, merge all overlapping intervals, and return an array of the non-overlapping intervals that cover all the intervals in the input. (Note: this is a classical interview question from Leetcode)

Example:

Input: [[1,3],[2,6],[8,10],[15,18]]

Output: [[1,6],[8,10],[15,18]]

```python
def merge_intervals(intervals):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
assert merge_intervals([[1,3],[2,6],[8,10],[15,18]]) == [[1,6],[8,10],[15,18]]
assert merge_intervals([[1,6],[2,4],[3,5]]) == [[1,6]]
assert merge_intervals([[-4,-1],[-2,3],[2,6]]) == [[-4,6]] 
```

### Question 3
Define a function called `switch_key_value`, which takes a dictionary as its only argument. Switch the placement of key and value of the dictionary and return it as a new dictionary.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
def switch_key_value(d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a0b5ecba31ff8afb, points: 1, grade: true, solution: false -->
```python
assert(switch_key_value({}) == {})
assert(switch_key_value({'a': 1, 'b': 2}) == {1: 'a', 2: 'b'})
```

### Question 4
Define a function called `sort_list_dict`, which takes a dictionary whose values are lists as its only argument. Sort all these lists alphabetically return them in a new dictionary (with the same keys).

<!-- grade_id: cell-996c61e79b68ca6d, points: 0, grade: false, solution: true -->
```python
def sort_list_dict(d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f752b519fbf2ac78, points: 1, grade: true, solution: false -->
```python
assert(sort_list_dict({}) == {})

assert(sort_list_dict({'a': []}) == {'a': []})

input_dict = {'a': [4, 1, 2, 3], 'b': [3, 2, 1]}
desired_output = {'a': [1, 2, 3, 4], 'b': [1, 2, 3]}
assert(sort_list_dict(input_dict) == desired_output)

input_dict = {'S': [2, 3, 1], 'T': [5, 1, 2], 'D': [3, 2, 4]}
desired_output = {'S': [1, 2, 3], 'T': [1, 2, 5], 'D': [2, 3, 4]}
assert(sort_list_dict(input_dict) == desired_output)
```